---
rule_files:
  - "status.rule"
evaluation_interval: 5m
tests:
  # ProcessRestarted
  - interval: 5m
    input_series:
      - series: process_start_time_seconds{job="mainnet"}
        values: 0 1 2 3 4 5
    alert_rule_test:
      # event should fire
      - eval_time: 15m
        alertname: ProcessRestarted
        exp_alerts:
          - exp_labels:
              severity: critical
              job: mainnet
            exp_annotations:
              summary: Too much AAVE mainnet bot restarts
              description: >-
                The AAVE mainnet bot has been restarted for 3 times
                for the last 15 minutes

  # StaleBotReport
  - interval: 5m
    input_series:
      - series: aave_bot_processing_finished_seconds{job="mainnet"}
        values: _x8 900+0x37
    alert_rule_test:
      - eval_time: 35m
        alertname: StaleBotReportOnLaunch
        exp_alerts:
          - exp_labels:
              severity: high
              job: mainnet
            exp_annotations:
              summary: Stale AAVE mainnet bot report detected
              description: >-
                No report has been received in the last 30
                minutes since the AAVE mainnet bot was launched
          - exp_labels:
              severity: high
              job: arbitrum
            exp_annotations:
              summary: Stale AAVE arbitrum bot report detected
              description: >-
                No report has been received in the last 30
                minutes since the AAVE arbitrum bot was launched
          - exp_labels:
              severity: high
              job: polygon
            exp_annotations:
              summary: Stale AAVE polygon bot report detected
              description: >-
                No report has been received in the last 30
                minutes since the AAVE polygon bot was launched

      - eval_time: 3h40m
        alertname: StaleBotReport
        exp_alerts:
          - exp_labels:
              severity: high
              job: mainnet
            exp_annotations:
              summary: Stale AAVE mainnet bot report detected
              description: Last report has been received more than 3 hours ago on AAVE mainnet bot

# vim: set ts=2 sw=2 ft=yaml:
